# Week 3 Exercises

See: _McKinney 2.3_ and [Python Documentation](https://docs.python.org/3/tutorial/controlflow.html) section 4 on flow control.


**At the begining of the semester, all of the workshop programming exercises will be structured a specific way to make it easier to verify for yourself that you're on the right track as well as easier for me to do a first pass on automated grading.  The structure of each question will require you to write a function using Python code. Don't worry that we haven't talked about functions yet. Just edit the code between** `### BEGIN SOLUTION` and `### END SOLUTION` **as shown in the example below.**


**WHAT I PROVIDE:**
```
def some_function(parameter1, parameter2):

   ### BEGIN SOLUTION
   x = -1
   ### END SOLUTION
   
   return x
```

**WHAT YOU SHOULD DO:** Just change the parameter names (if you feel you need to) and the calculations between `### BEGIN SOLUTION` and `### END SOLUTION`.  This is just a made up example.
```
def some_function(a, b):

   ### BEGIN SOLUTION
   temp = a + b
   x = temp / a * b
   ### END SOLUTION
   
   return x
```


---
---

**Below each programming exercise are some tests (`assertions`) that verify your code is working correctly.  If any assertions fail, you know that something isnt' right with your code, but having all assertions pass doesn't necessarily mean your code is perfect, yet. You should also create your own tests to make sure your code is correct.**

**For now, please don't change any function names**

### 14.1 Difference in rate per 1,000

Often in public health, we report metrics as a number per 1,000 or per 1,000,000 in population. The purpose of that is to normalize the numbers between area of larger and smaller populations.

Below, we have a function already built to calculate teh different in rate per 1,000 in two different regions.  The function normalizes them and then prints out a message describing their difference.

```python
def diff_in_rate_normalized(count_AP, total_AP, count_TG, total_TG):
    """(int,int,int,int) -> str
    * count_a is the number of occurences in area A
    * total_a is the total population in area A
    * count_b is the number of occurences in area B
    * total_b is the total population in area B
    
    This function returns a string describing how A and B compare in terms of occurence rates per 1,000 population."""

    rate_AP = count_AP / total_AP * 100000
    rate_TG = count_TG / total_TG * 100000
    
    
    if rate_AP == rate_TG:
        msg = "The rate in AP and the rate in TG are the same ({}).".format(rate_AP)
    elif rate_AP >= rate_TG:
        msg = "The rate in AP ({}) is slightly higher than the rate in TG ({}).".format(rate_AP, rate_TG)
    elif rate_AP <= rate_TG:
        msg = "The rate in AP ({}) is slightly less than the rate in TG ({})".format(rate_AP, rate_TG)
    else:
        msg = "The rate in AP ({}) is not equal to rate in TG ({})".format(rate_AP,rate_TG)
        
    return msg
```

```python
rate_covid_testing = diff_in_rate_normalized(20,10000,18,10000)
print(rate_covid_testing)
```

Output:
```
The rate in AP (200.0) is slightly higher than the rate in TG (180.0).
```

```python
diff_in_rate_normalized(3,10000,5,10000)
```

Output:
```
'The rate in AP (29.999999999999996) is slightly less than the rate in TG (50.0)'
```

### 14.2 Trimming outliers

While it works well in most cases, in practice, this may not truly make numbers comparable between extremely large population centers (e.g. New York City at 8.5 million) and very small rural areas (e.g. Meeteese, WY at 459 people).

Let's take that function and make some adjustments.  If the total population of either A or B is more than 100 times larger than the other, then we want to return a different message.  That is, if the populations are more than two orders of magnitude different, then we shouldn't try to compare them.

```python
def diff_in_rate_normalized(count_AP, total_AP, count_TG, total_TG):
    """(int,int,int,int) -> str
    * count_a is the number of occurences in area A
    * total_a is the total population in area A
    * count_b is the number of occurences in area B
    * total_b is the total population in area B
    
    This function returns a string describing how A and B compare in terms of occurence rates per 1,000 population.
    
    If total_a / total_b > 100 OR total_b / total_a > 100 then we'll return a message saying the two can't be compared.
    "The total populations in A and B are so different that they can't be compared."
    """
    
    msg = ""
    
    if ((total_AP / total_TG) > 100) or ((total_TG / total_AP) > 100):
        msg = "The total populations in AP and TG are so different that they can't be compared."
    else:
        rate_AP = count_AP / total_AP * 100000
        rate_TG = count_TG / total_TG * 100000

        if rate_AP == rate_TG:
        msg = "The rate in AP and the rate in TG are the same ({}).".format(rate_AP)
        elif rate_AP >= rate_TG:
        msg = "The rate in AP ({}) is slightly higher than the rate in TG ({}).".format(rate_AP, rate_TG)
        elif rate_AP <= rate_TG:
        msg = "The rate in AP ({}) is slightly less than the rate in TG ({}).".format(rate_AP, rate_TG)
        else:
        msg = "The rate in AP ({}) is not equal to rate in TG ({}).".format(rate_AP,rate_TG)
        
    return msg
```

```python
diff_in_rate_normalized(5, 459, 30, 8500)
```

Output:
```
'The rate in AP (1089.3246187363834) is slightly higher than the rate in TG (352.94117647058823).'
```

```python
rate_covid_testing = diff_in_rate_normalized(21,10000,19,10000)
print(rate_covid_testing)
```

Output:
```
The rate in AP (210.0) is slightly higher than the rate in TG (190.0).
```

### 14.3 LACE Score
There is a simple readmission index called the LACE Score: https://www.hindawi.com/journals/bmri/2015/169870/tab1/

Use this documentation to create a function that can compute LACE scores based on the 4 input parameters.  Below is the function signature and documentation that you start with.

<!-- grade_id: bmi-answer, points: 0, grade: false, solution: true -->
```python
def LACE(length_of_stay, acute_flag, charlson_index, ed_visits):
    """(int, bool, int, int) -> int
    This function uses the logic from https://www.hindawi.com/journals/bmri/2015/169870/tab1/
    to compute the LACE score for this patient.
    
    >>> LACE(4, False, 1, 0)
    5
    
    >>> LACE(4, True, 4, 7)
    16
    
    """
 
    ### BEGIN SOLUTION
    if length_of_stay < 1: length_of_stay=0
    elif length_of_stay==1: length_of_stay=1
    elif length_of_stay==2: length_of_stay=2
    elif length_of_stay==3: length_of_stay=3
    elif length_of_stay>=4 and length_of_stay<=6 : length_of_stay=4
    elif length_of_stay>=7 and length_of_stay<=13 : length_of_stay=5
    elif length_of_stay>=14: length_of_stay=7    
    
    if acute_flag==True: acute_flag=3
    elif acute_flag==False: acute_flag=0
    
    if charlson_index==0: charlson_index=0
    elif charlson_index==1: charlson_index=1
    elif charlson_index==2: charlson_index=2
    elif charlson_index==3: charlson_index=3
    elif charlson_index>=4: charlson_index=5
        
    if ed_visits==0: ed_visits=0
    elif ed_visits==1: ed_visits=1
    elif ed_visits==2: ed_visits=2
    elif ed_visits==3: ed_visits=3
    elif ed_visits>=4: ed_visits=4
        
    score=length_of_stay+acute_flag+charlson_index+ed_visits
    
    ### END SOLUTION
    
    
    return score
```

<!-- grade_id: bmi-test, points: 10, grade: true, solution: false -->
```python
assert LACE(4, False, 1, 0) == 5
assert LACE(4, True, 4, 7) == 16
```

```python
### 14.4 Care Management Criteria

Care managers use LACE as part of the criteria for assigning a care coodinator to a patient who has been recently discharged. If the score is above 10, then a care coordinator will be assigned. The other criteria they use is if the patient has been discharged with a diagnosis of CHF or COPD.  If the diagnosis field has CHF or COPD in it, then the patient will have a care coordinator assigned.

For this exercise, write another function that takes the same inputs as LACE() plus another diagnosis parameter, and return True or False depending on if the paient needs a care coordinator.

**NOTE** Pay attention to the fact that the order of parameters in this function definition are not the same as the order they were in the LACE score.  
```

<!-- grade_id: weight-answer, points: 0, grade: false, solution: true -->
```python
def assign_care_coordinator(diagnosis_cd, ed_visits, length_of_stay, acute_flag, charlson_index):
    """ (str, int, int, bool, int) -> bool
    Care managers use LACE as part of the criteria for assigning a care coodinator to a 
    patient who has been recently discharged. If the score is above 10, then a care 
    coordinator will be assigned. The other criteria they use is if the patient has been 
    discharged with a diagnosis of CHF or COPD. If the diagnosis field has CHF or COPD in 
    it, then the patient will have a care coordinator assigned.
    """
    
    ### BEGIN SOLUTION
    if diagnosis_cd=='CHF' or diagnosis_cd=='COPD': assignment=True
    else:
        
        if ed_visits==0: ed_visits=0
        elif ed_visits==1: ed_visits=1
        elif ed_visits==2: ed_visits=2
        elif ed_visits==3: ed_visits=3
        elif ed_visits>=4: ed_visits=4
    
        if length_of_stay < 1: length_of_stay=0
        elif length_of_stay==1: length_of_stay=1
        elif length_of_stay==2: length_of_stay=2
        elif length_of_stay==3: length_of_stay=3
        elif length_of_stay>=4 and length_of_stay<=6 : length_of_stay=4
        elif length_of_stay>=7 and length_of_stay<=13 : length_of_stay=5
        elif length_of_stay>=14: length_of_stay=7    

        if acute_flag==True: acute_flag=3
        elif acute_flag==False: acute_flag=0

        if charlson_index==0: charlson_index=0
        elif charlson_index==1: charlson_index=1
        elif charlson_index==2: charlson_index=2
        elif charlson_index==3: charlson_index=3
        elif charlson_index>=4: charlson_index=5

        score=length_of_stay+acute_flag+charlson_index+ed_visits

        if score>10: assignment=True
        else: assignment=False

    
    ### END SOLUTION
        
    return assignment
    
```

<!-- grade_id: weight-tests, points: 5, grade: true, solution: false -->
```python
assert assign_care_coordinator('None', 0, 4, False, 1) == False
assert assign_care_coordinator('CHF', 0, 4, False, 1) == True
assert assign_care_coordinator('COPD', 0, 4, False, 1) == True
assert assign_care_coordinator('None', 7, 4, True, 4) == True
assert assign_care_coordinator('CHF', 7, 4, True, 4) == True
```

```python
assign_care_coordinator('None', 0, 4, False, 1)
```

### 14.5 qCSI COVID-19 Severity Index

See: https://www.mdcalc.com/quick-covid-19-severity-index-qcsi#evidence

Calculate the total risk score as per the point values assigned to respiratory rate, pulse oximetry, and O2 flow rate.  Then calculate and return the Risk Level.

In addition to the rules provided at the link above, also add the following checks for valid values:
* If `respiratory_rate <= 0` then return _invalid respiratory rate_
* If `pulse_ox <= 0` then return _invalid pulse ox_
* If `pulse_ox > 100` then return _invalid pulse ox_
* If `os_flow <= 0` then return _invalid O2 flow rate_

```python
def qcsi(respiratory_rate, pulse_ox, o2_flow):
    """(int, int, int) -> str
    * respiratory_rate is an integer value
    * pulse_ox is an integer value (e.g. 30 means 30%)
    * o2_flow is an integer value
    """
    
    ### BEGIN SOLUTION
    
    if respiratory_rate<=22: respiratory_rate=0
    elif respiratory_rate>=23 and respiratory_rate<=28: respiratory_rate=1
    elif respiratory_rate>28: respiratory_rate=2
    
    if  pulse_ox>92:  pulse_ox=0
    elif  pulse_ox>=89 and  pulse_ox<=92:  pulse_ox=2
    elif  pulse_ox<=88:  pulse_ox=5
        
    elif o2_flow<=2: o2_flow=0
    elif o2_flow==3 or o2_flow==4: o2_flow=4
    elif o2_flow==5 or o2_flow==6: o2_flow=5
    
    score = respiratory_rate+pulse_ox+o2_flow
    
    if score<=3: risk='low'
    elif score>=4 and score<=6: risk='low-intermediate'
    elif score>=7 and score<=9: risk='high-intermediate'
    elif score>=10 and score<=12: risk='high'
    
    if respiratory_rate<=0: risk='invalid respiratory rate'
    if  pulse_ox<=0 or  pulse_ox > 100 : risk='invalid pulse ox'
    if o2_flow<=0: risk='invalid flow rate'
    
        
            
    ### END SOLUTION
    
    return risk
```

```python
assert (qcsi(29, 95, 1) == 'low')
assert (qcsi(20, 93, 1) == 'low')
assert (qcsi(29, 88, 1) == 'high-intermediate')
assert (qcsi(29, 88, 4) == 'high')
assert (qcsi(30, 90, 1) == 'low-intermediate')
assert (qcsi(28, 92, -1) == 'invalid O2 flow rate')
assert (qcsi(22, 0, 4) == 'invalid pulse ox')
assert (qcsi(0, 97, 2) == 'invalid respiratory rate')
```

---

## Check your work above

If you didn't get them all correct, take a few minutes to think through those that aren't correct.


## Submitting Your Work

In order to submit your work, you'll need to use the `git` command line program to **add** your homework file (this file) to your local repository, **commit** your changes to your local repository, and then **push** those changes up to github.com.  From there, I'll be able to **pull** the changes down and do my grading.  I'll provide some feedback, **commit** and **push** my comments back to you.  Next week, I'll show you how to **pull** down my comments.

First run through everything one last time and submit your work:
1. Use the `Kernel` -> `Restart Kernel and Run All Cells` menu option to run everything from top to bottom and stop here.
2. Then open a new command line by clicking the `+` icon above the file list and chosing `Terminal`
3. At the command line in the new Terminal, follow these steps:
  1. Change directories to your project folder and the week03 subfolder (`cd <folder name>`)
  2. Make sure your project folders are up to date with github.com (`git pull`)
  3. Add the homework files for this week (`git add <file name>`)
  4. Commit your changes (`git commit -m "message"`)
  5. Push your changes (`git push`)
  
**Here's a full examle**
```
cd hds5210-2022/week03
git pull
git add week03_assignment_1.ipynb week03_assignment_2.ipynb
git commit -a -m "Submitting homework assignments for week 2"
git push
```

If anything fails along the way with this submission part of the process, let me know.  I'll help you troubleshoort.